Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Landyn Francis"
COLLABORATORS = ""
```

---

# Exercise 1 
(10 marks)

$\newcommand{\bfx}{\mathbf{x}}$
$\newcommand{\bfb}{\mathbf{b}}$
Find the minimum of the following function
$$f(\bfx) = \bfx^\top A^\top A \bfx + 2 \bfb^\top \bfx + c$$

To find the minimum of $f(\bfx)$, we must take the derivative, and set it equal to zero. 
$$f'(\bfx) = 2A^\top A \bfx + 2\bfb = 0$$
Moving $\bfb$ to the other side, we get the following:
$$-\bfb=A^\top A\bfx$$
Because we cannot divide matrices, we must multiply by the inverse product to isolate $\bfx$
$$\bfx=-(A^\top A)^{-1}\bfb$$

To determine the minimum value, we must plug in $\bfx$ into the original equation. Looking ahead, we will first caluclate $\bfx^\top$, as it must be plugged in as well.
First, apply the multiplication transpose rule.
$$\bfx^\top = -\bfb^\top((A^\top A)^{-1})^{T}$$
Then, apply the inner inverse and transpose.
$$\bfx^\top = -\bfb^\top A^{-1}(A^\top)^{-1}$$
Plugging these values into $f(\bfx)$ results in the following:
$$f(\bfx)|_{\bfx} = (-\bfb^\top A^{-1} (A^\top)^{-1}(A^\top A)(-(A^\top A)^{-1}\bfb)+2\bfb^\top (-(A^\top A)^{-1}\bfb)+c$$
After cancelling terms and simplifying, we get the following minimum value:
$$f(\bfx)_{min} = -\bfb^\top (A^\top A)^{-1}\bfb+c$$


## Exercise 2
(10 marks)

$\newcommand{\bfmu}{\boldsymbol{\mu}}$

$\newcommand{\x}{\mathbf{x}}$
$\newcommand{\muone}{\boldsymbol{\mu_1}}$
$\newcommand{\mutwo}{\boldsymbol{\mu_2}}$
$\newcommand{\sigmaone}{\boldsymbol{\Sigma_1}}$
$\newcommand{\sigmatwo}{\boldsymbol{\Sigma_2}}$
$\newcommand{\inv}{^{-1}}$
$\newcommand{\sigsum}{(\sigmaone+\sigmatwo)}$
$\newcommand{\sigmusum}{(\sigmaone\muone+\sigmatwo\mutwo)}$
$\newcommand{\tp}{^\top}$
$\newcommand{\xthalf}{(\muone\tp\sigmaone+\mutwo\tp\sigmatwo)}$

Find the minimum of the following function (assume $\Sigma_1$ and $\Sigma_2$ to be positive semi-definite).
$$f(\bfx) = (\bfx - \bfmu_1)^\top \Sigma_1 (\bfx - \bfmu_1)
+ (\bfx - \bfmu_2)^\top \Sigma_2 (\bfx - \bfmu_2) $$

As with Exercise 1, to find the minimum we must first take the derivative, then set it equal to 0.
If we define $\mathbf{y} = \x-\muone$, then the following is clear:
$$\mathbf{y}^\top \sigmaone \mathbf{y} = 2\sigmaone \mathbf{y}$$
With this we can replace $\mathbf{y}$ with $(\x-\muone)$, find $f'(\x)$, and set it equal to 0.
$$f'(\x) = 2\sigmaone(\x-\muone)+2\sigmatwo(\x-\muone) = 0$$
Dividing by 2, and distributing the $\sigmaone$ and $\sigmatwo$ through the equation, we get:
$$0 = \sigmaone\x - \sigmaone\muone+\sigmatwo\x-\sigmatwo\mutwo$$
Grouping like terms:
$$\sigmaone\muone+\sigmatwo\mutwo = \sigsum\x$$
Isolating x, we find the $\x$ leading to the minimum value.
$$\x = \sigsum\inv \sigmusum$$
Again, we must find $\x\tp$, as it must be plugged in as well. Applying the transpose rule we get the following:
$$\x\tp = \sigmusum\tp (\sigsum\tp)\inv$$
Because $\sigmaone$ and $\sigmatwo$ are symmetric positive semi-definite, we know that $\sigmaone = \sigmaone\tp$ $\sigmatwo = \sigmatwo\tp$. With this we can apply the addition transpose rule on each term.
$$\x\tp = ((\sigmaone\muone)\tp + (\sigmatwo\mutwo)\tp) (\sigmaone\tp+\sigmatwo\tp)\inv$$
Finally, apply the multiplication transpose rule.
$$\x\tp = (\muone\tp\sigmaone+\mutwo\tp\sigmatwo)\sigsum\inv$$
Next, before we plug in $\x$ and $\x\tp$ we must expand and factor $f(\x)$, it expands to
$$(\x\tp\sigmaone\x) - (\x\tp\sigmaone\muone) - (\muone\tp\sigmaone\x) + (\muone\tp\sigmaone\muone)$$ 
$$ + $$
$$(\x\tp\sigmatwo\x) - (\x\tp\sigmatwo\mutwo) - (\mutwo\tp\sigmatwo\x) + (\mutwo\tp\sigmatwo\mutwo)$$
The equation has been arranged so the factors are easier to see. Factoring out terms results in the simplified $f(\x)$
$$f(\x) = (\x\tp\sigsum\x) - \x\tp\sigmusum - (\muone\tp\sigmaone+\mutwo\tp\sigmatwo)\x + 
(\muone\tp\sigmaone\muone) + (\mutwo\tp\sigmatwo\mutwo)$$
Now we are ready to plug in $\x$ and $\x\tp$, for clarity, each term will be placed on it's own line, and denoted by $t_n$
$$f(\x)|_{\x} = $$
$$t_1 = +(\xthalf\sigsum\inv\sigsum\sigsum\inv\sigmusum)$$
$t_1$ can be simplified to
$$t_1 = +\xthalf\sigsum\inv\sigmusum$$
$$t_2 = -\xthalf\sigsum\inv\sigmusum$$
$$t_3 = -(\muone\tp\sigmaone+\mutwo\tp\sigmatwo)\sigsum\inv\sigmusum $$
$$t_4 = +(\muone\tp\sigmaone\muone)$$ 
$$t_5 = +(\mutwo\tp\sigmatwo\mutwo)$$
It can be seen that $t_1,t_2$, and $t_3$ are equivalent, and therefore $a-a-a=-a$, and $t_4$, and $t_5$ cannot be simplified. This results in the final minimum value being equivalent to.
$$min f(\bfx) = -\xthalf\sigsum\inv\sigmusum+(\muone\tp\sigmaone\muone)+(\mutwo\tp\sigmatwo\mutwo)$$

## Exercise 3: Fit a plane to point cloud

$\newcommand{\x}{\mathbf{x}}$
$\newcommand{\y}{\mathbf{y}}$
$\newcommand{\z}{\mathbf{z}}$
$\newcommand{\bzeta}{\boldsymbol{\zeta}}$




Least square fit in 2D fits a line to a given set of 2D points. This homework asks you to repeat the minimization procedure to find a plane that fits a 2D points. Below you are provided with helper code to load a point cloud, and visualize the points and the plane. Your task is to find the equation of plane that fits the point cloud. 

1. Show the math and derivation for your code

The easiest way to find the plane of best fit for a set of 3D points is to vectorize the equation and solve for the unknown values. Our base equation is as follows:
$$z = ax+by+c$$
We are given data in the form of 3D points with an x, y, and z value. 
We can vectorize this data, by creating vectors $\x,\y$, and $\z$ such that the values of the vectors are equal to each x,y, and z point provided.
$$\x = \begin{bmatrix} x_1 \\ x_2 \\ ... \\ x_n\end{bmatrix}, \y = \begin{bmatrix} y_1 \\ y_2 \\ ... \\ y_n\end{bmatrix}, \z = \begin{bmatrix} z_1 \\ z_2 \\ ... \\ z_n\end{bmatrix}$$
We can also define a column vector $\bzeta$ for our vector of unknown values.
$$\bzeta = \begin{bmatrix} a \\ b \\ c \end{bmatrix}$$
Next, we want our equation in the form $A\mathbf{x}=\mathbf{b}$, to do this, we can pack a matrix, $\boldsymbol{\Delta}$, with the $\x$ and $\y$ vectors, set $\mathbf{x}$ equal to our $\bzeta$ vector, and use $\z$ as $\mathbf{b}$. However, this would be multiplying an nx2 matrix with a 3x1 vector, which is not possible. To work around this, we must introduce a column of ones into $\boldsymbol{\Delta}$. This results in the following vectorized equation.
$$\boldsymbol{\Delta}\bzeta=\z$$
$$\begin{bmatrix}x & y & (\mathbb{1}_n)^\top \end{bmatrix} \begin{bmatrix} a \\ b \\ c \end{bmatrix} = \z$$
$$\begin{bmatrix}x_1 & y_1 & 1 \\
                 x_2 & y_2 & 1 \\
                 ... & ... & ... \\
                 x_n & y_n & 1 \\ \end{bmatrix} \begin{bmatrix} a \\ b \\ c \end{bmatrix} = \begin{bmatrix} z_1 \\ z_2 \\ ... \\ z_n \end{bmatrix}$$
Next we can solve for $\bzeta$, the vector containing the coefficients a,b,c, that we need. Normally, we would take the inverse of $\boldsymbol{\Delta}$. However there are two reasons why the pseudo-inverse must be used. The first reason is that $\boldsymbol{\Delta}$ is not a square matrix, and therefore it is non-invertible. The second reason is the system is over-determined, which means we have more data points than is ultimately necessary to find the equation of a plane. Therefore, the pseudo inverse must be used. The pseudo inverse is defined as, $A^\dagger = (A^\top A)^{-1}A^\top$. Using this we can solve for $\bzeta$.
$$\bzeta = \boldsymbol{\Delta}^\dagger \z$$
$\bzeta$ is the vector of coefficients of the best fitting plane, so the final equation is as follows:
$$\begin{bmatrix} a \\ b \\ c \end{bmatrix} = \boldsymbol{\Delta}^\dagger \z$$
2. Implement the code in python and numpy to find the coefficients of the equation of plane and write your answer in the space provided below. Submit this ipynb file to brightspace. (30 marks)

```python
#Comment after running this once
#!pip install --user k3d


import k3d
from k3d import matplotlib_color_maps

def k3d_init_plot():
    return k3d.plot(grid_visible=False,
                    camera_auto_fit=True,
                    camera_fov=30.0,
                    background_color=0xffffff)

def visualize_plane(pts, abc, k3dplt=None):
    _k3dplt = k3dplt
    k3dplt = k3d_init_plot() if k3dplt is None else k3dplt
    a, b, c = abc
    x = np.linspace(-1, 1, 100, dtype=np.float32)
    y = np.linspace(0, 1, 100, dtype=np.float32)

    x, y = np.meshgrid(x, y)
    z = a*x + b*y + c
    plt_surface = k3d.surface(z,
                              xmin=-1, xmax=1,
                              ymin=0, ymax=1,
                              compression_level=9,
                              wireframe=True,
                              color_map=matplotlib_color_maps.Coolwarm_r,
                              attribute=z, color_range=[-1, 0.5])
    
    k3dplt += plt_surface
    visualize_points(pts, k3dplt=k3dplt)
    if _k3dplt is None:
        k3dplt.display()

def visualize_points(pts, k3dplt=None):
    _k3dplt = k3dplt
    k3dplt = k3d_init_plot() if k3dplt is None else k3dplt
    plt_points = k3d.points(pts,
                            attribute=pts[:, 2], # use z coordinate to color
                            point_size=0.005,
                            shader="flat",
                            color_map=matplotlib_color_maps.Inferno)
    k3dplt += plt_points
    if _k3dplt is None:
        k3dplt.display()
```

```python
import numpy as np
#!wget https://github.com/wecacuee/ECE490-Neural-Networks/raw/master/notebooks/02-linear-models/table_top.npy
pts = np.load('table_top.npy').astype('float32') # Nx3
visualize_points(pts)
pts.min(axis=0), pts.max(axis=0)
```

Output:
```
Output()
(array([-1.0558e+00,  3.2414e-04, -1.8675e+00], dtype=float32),
 array([ 0.59304,  0.53329, -1.0252 ], dtype=float32))
```

<!-- grade_id: cell-bb7ef5a91325977a, points: 0, grade: false, solution: true -->
```python
def plane_fit(pts):
    #Make a column vector out of the z values
    b = pts[:,-1]
    #Pack a matrix full of a column of x values, a column of y values, and a column of 1's
    A = np.column_stack((pts[:,0],pts[:,1],np.ones_like(b)))
    # abc = A(pseudo-inverse) multiplied by vector b.
    abc = np.linalg.pinv(A)@b
    return abc
```

<!-- grade_id: cell-5b7358c1e61a4706, points: 50, grade: true, solution: false -->
```python
# Find abc such that z = a x + b y + c
abc = plane_fit(pts)
visualize_plane(pts, abc)
assert np.abs(pts[:,2] - np.hstack((pts[:, :2], np.ones((pts.shape[0],1)))) @ abc).mean() < 0.4

```

Output:
```
Output()
```