# Exercise 1 
(10 marks)

$\newcommand{\bfx}{\mathbf{x}}
\newcommand{\bfb}{\mathbf{b}}$
Find the minimum of the following function
$$f(\bfx) = \bfx^\top A^\top A \bfx + 2 \bfb^\top \bfx + c$$

<!-- grade_id: cell-be3d981bec8a49f9, points: 10, grade: true, solution: true -->

At the minimum point $\bfx^*$, the vector derivative is zero

$\newcommand{\p}{\partial }$
$\newcommand{\bfzero}{\mathbf{0}}$
$$ \frac{\p}{\p \bfx} f(\bfx) |_{\bfx^*} = \bfzero^\top$$
$$\implies 2{\bfx^*}^\top (A^\top A) + 2\bfb^\top =  \bfzero^\top$$
$$\implies 2 (A^\top A)^\top\bfx^* + 2\bfb =  \bfzero$$
$$\implies 2 (A^\top A)\bfx^* + 2\bfb =  \bfzero$$
$$\implies \bfx^* = - (A^\top A)^{-1}\bfb $$

The minimimum value is
$$ f(\bfx^*) = \bfb^\top (A^\top A)^{-1}\bfb - 2 \bfb^\top (A^\top A)^{-1}\bfb + c$$ 
$$ f(\bfx^*) = - \bfb^\top (A^\top A)^{-1}\bfb + c$$ 


## Exercise 2
(10 marks)

$\newcommand{\bfmu}{\boldsymbol{\mu}}$
Find the minimum of the following function (assume $\Sigma_1$ and $\Sigma_2$ to be symmetric and positive semi-definite).
$$f(\bfx) = (\bfx - \bfmu_1)^\top \Sigma_1 (\bfx - \bfmu_1)
+ (\bfx - \bfmu_2)^\top \Sigma_2 (\bfx - \bfmu_2) $$

<!-- grade_id: cell-cb331ee9d8b9f9d8, points: 10, grade: true, solution: true -->

At the minimum point $\bfx^*$, the vector derivative is zero

$$ \frac{\p}{\p \bfx} f(\bfx) |_{\bfx^*} = \bfzero^\top$$
$$\implies 2 (\bfx^* - \bfmu_1)^\top (\Sigma_1^\top+\Sigma_1) + 2 (\bfx^* - \bfmu_2)^\top (\Sigma_2^\top+\Sigma_2) =  \bfzero^\top$$

$$\implies 2 (\bfx^* - \bfmu_1)^\top \Sigma_1 + 2 (\bfx^* - \bfmu_2)^\top \Sigma_2 =  \bfzero^\top$$

Group the $\bfx^*$ terms and the constant terms
$$\implies 2 {\bfx^*}^\top (\Sigma_1 + \Sigma_2) =   2 \bfmu_1^\top \Sigma_1 + 2\bfmu_2^\top \Sigma_2 $$

Take the transpose and positive definite matrices are always invertible (use $\Sigma_1^\top = \Sigma_1$ and $\Sigma_2^\top = \Sigma_2$):
$$\implies \bfx^* = (\Sigma_1 + \Sigma_2)^{-1} (\Sigma_1 \bfmu_1 + \Sigma_2 \bfmu_2)$$


To get the minimum value, group the quadratic terms, liner terms and constant togther.
$$f(\bfx^*) = {\bfx^*}^\top (\Sigma_1 + \Sigma_2) \bfx - 2 (\bfmu_1^\top  \Sigma_1 + \bfmu_2 \Sigma_2) \bfx
+ \bfmu_1^\top \Sigma_1 \bfmu_1 + \bfmu_2^\top \Sigma_2 \bfmu_2 $$

The minimum value is

$$f(\bfx^*) = - (\bfmu_1^\top  \Sigma_1 + \bfmu_2 \Sigma_2) (\Sigma_1 + \Sigma_2)^{-1} (\Sigma_1^\top \bfmu_1 + \Sigma_2^\top \bfmu_2)
+ \bfmu_1^\top \Sigma_1 \bfmu_1 + \bfmu_2^\top \Sigma_2 \bfmu_2 $$


## Exercise 3: Fit a plane to point cloud

Least square fit in 2D fits a line to a given set of 2D points. This homework asks you to repeat the minimization procedure to find a plane that fits a 2D points. Below you are provided with helper code to load a point cloud, and visualize the points and the plane. Your task is to find the equation of plane that fits the point cloud. 

1. Show the math and derivation for your code
2. Implement the code in python and numpy to find the coefficients of the equation of plane and write your answer in the space provided below. Submit this ipynb file to brightspace. (30 marks)

```python
#Comment after running this once
!pip install --user k3d


import k3d
from k3d import matplotlib_color_maps

def k3d_init_plot():
    return k3d.plot(grid_visible=False,
                    camera_auto_fit=True,
                    camera_fov=30.0,
                    background_color=0xffffff)

def visualize_plane(pts, abc, k3dplt=None):
    _k3dplt = k3dplt
    k3dplt = k3d_init_plot() if k3dplt is None else k3dplt
    a, b, c = abc
    x = np.linspace(-1, 1, 100, dtype=np.float32)
    y = np.linspace(0, 1, 100, dtype=np.float32)

    x, y = np.meshgrid(x, y)
    z = a*x + b*y + c
    plt_surface = k3d.surface(z,
                              xmin=-1, xmax=1,
                              ymin=0, ymax=1,
                              compression_level=9,
                              wireframe=True,
                              color_map=matplotlib_color_maps.Coolwarm_r,
                              attribute=z, color_range=[-1, 0.5])
    
    k3dplt += plt_surface
    visualize_points(pts, k3dplt=k3dplt)
    if _k3dplt is None:
        k3dplt.display()

def visualize_points(pts, k3dplt=None):
    _k3dplt = k3dplt
    k3dplt = k3d_init_plot() if k3dplt is None else k3dplt
    plt_points = k3d.points(pts,
                            attribute=pts[:, 2], # use z coordinate to color
                            point_size=0.005,
                            shader="flat",
                            color_map=matplotlib_color_maps.Inferno)
    k3dplt += plt_points
    if _k3dplt is None:
        k3dplt.display()
```

Output:
```
Requirement already satisfied: k3d in /home/vdhiman/.local/lib/python3.11/site-packages (2.15.2)
Requirement already satisfied: ipywidgets<9.0.0,>=7.0.0 in /home/vdhiman/.local/miniconda3/envs/nbgrader/lib/python3.11/site-packages (from k3d) (8.0.4)
Requirement already satisfied: msgpack in /home/vdhiman/.local/lib/python3.11/site-packages (from k3d) (1.0.4)
Requirement already satisfied: numpy in /home/vdhiman/.local/miniconda3/envs/nbgrader/lib/python3.11/site-packages (from k3d) (1.24.1)
Requirement already satisfied: traitlets in /home/vdhiman/.local/lib/python3.11/site-packages (from k3d) (5.9.0)
Requirement already satisfied: traittypes in /home/vdhiman/.local/lib/python3.11/site-packages (from k3d) (0.2.1)
Requirement already satisfied: ipykernel>=4.5.1 in /home/vdhiman/.local/miniconda3/envs/nbgrader/lib/python3.11/site-packages (from ipywidgets<9.0.0,>=7.0.0->k3d) (6.20.1)
Requirement already satisfied: ipython>=6.1.0 in /home/vdhiman/.local/miniconda3/envs/nbgrader/lib/python
```

```python
import numpy as np
!wget https://github.com/wecacuee/ECE490-Neural-Networks/raw/master/notebooks/02-linear-models/table_top.npy
pts = np.load('table_top.npy').astype('float32') # Nx3
visualize_points(pts)
pts.min(axis=0), pts.max(axis=0)
```

Output:
```
--2023-02-02 02:45:49--  https://github.com/wecacuee/ECE490-Neural-Networks/raw/master/notebooks/02-linear-models/table_top.npy
Resolving github.com (github.com)... 140.82.113.3
Connecting to github.com (github.com)|140.82.113.3|:443... connected.
HTTP request sent, awaiting response... 302 Found
Location: https://raw.githubusercontent.com/wecacuee/ECE490-Neural-Networks/master/notebooks/02-linear-models/table_top.npy [following]
--2023-02-02 02:45:49--  https://raw.githubusercontent.com/wecacuee/ECE490-Neural-Networks/master/notebooks/02-linear-models/table_top.npy
Resolving raw.githubusercontent.com (raw.githubusercontent.com)... 2606:50c0:8002::154, 2606:50c0:8000::154, 2606:50c0:8001::154, ...
Connecting to raw.githubusercontent.com (raw.githubusercontent.com)|2606:50c0:8002::154|:443... connected.
HTTP request sent, awaiting response... 200 OK
Length: 3331712 (3.2M) [application/octet-stream]
Saving to: ‘table_top.npy.1’

table_top.npy.1     100%[===================>]   3.18M  11.
```

<!-- grade_id: cell-bb7ef5a91325977a, points: 0, grade: false, solution: true -->
```python
def plane_fit(pts):
    ### BEGIN SOLUTION
    # z = a x + b y + c
    n = pts.shape[0]
    bfz = pts[:, 2]
    bfx = pts[:, :2]
    bfX = np.hstack((pts[:, :2], np.ones((n,1))))
    abc = np.linalg.inv(bfX.T @ bfX) @ (bfX.T @ bfz)
    ### END SOLUTION
    return abc
```

<!-- grade_id: cell-5b7358c1e61a4706, points: 50, grade: true, solution: false -->
```python
# Find abc such that z = a x + b y + c
abc = plane_fit(pts)
visualize_plane(pts, abc)
assert np.abs(pts[:,2] - np.hstack((pts[:, :2], np.ones((pts.shape[0],1)))) @ abc).mean() < 0.4

```

Output:
```
Output()
```